# Week 11 Problem 2

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do not write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select *Kernel*, and restart the kernel and run all cells (*Restart & Run all*).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select *File* → *Save and CheckPoint*)

5. When you are ready to submit your assignment, go to *Dashboard* → *Assignments* and click the *Submit* button. Your work is not submitted until you click *Submit*.

6. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

7. **If your code does not pass the unit tests, it will not pass the autograder.**

# Due Date: 6 PM, April 09, 2018

## Author: Kelechi Ikegwu
### Primary Reviewer: John Nguyen

```python
from nose.tools import assert_equal, assert_true, assert_false, assert_almost_equal
import numpy.testing as npt
import tweepy as tw

```

### Problem 2.1

Complete the function *connect* by first creating a twitter.cred file in your home directory (*/home/data_scientist/*). Read in the tokens and secrets from the twitter.cred file. Pass in the consumer token and secret using tweepy's [OAuthHandler](http://docs.tweepy.org/en/v3.5.0/auth_tutorial.html) function and set the access token and secret to  the access token and access sceret from your twitter.cred file. Lastly pass in the *OAuthHandler* object to tweepy's [Twitter API wrapper](http://docs.tweepy.org/en/v3.5.0/api.html) constructor and return it.


<!-- grade_id: p21_ans, points: 0, grade: false, solution: true -->
```python
def connect(path):
    '''
    Parameters
    ----------

    path: string containing a path to a twitter.cred file.
    
    Returns
    -------

    A Twitter API wrapper of type tweppy.api.API
    '''
    # YOUR CODE HERE
    tokens = []
    with open(path, 'r') as fin:
        for line in fin:
            if line[0] != '#':
                tokens.append(line.rstrip('\n'))
                
    auth = tw.OAuthHandler(tokens[2],tokens[3])
    auth.set_access_token(tokens[0],tokens[1])
    
    api = tw.API(auth, wait_on_rate_limit=True)
    
    return api
```

<!-- grade_id: p21_tests, points: 10, grade: true, solution: false -->
```python
path = '/home/data_scientist/twitter.cred'
api = connect(path)
assert_true(type(api) == type(tw.api))
try:
    api.me()
except tw.TweepError as e:
    print(e)
    assert_true(False, msg=e)
```

### Problem 2.2

Complete the function *get_screen_name* by using the tweepy's api wrapper to get the screen name of the twitter id passed in. 

<!-- grade_id: p22_ans, points: 0, grade: false, solution: true -->
```python
def get_screen_name(_api, _id):
    '''
    Parameters
    ----------
    _api: A Twitter API wrapper of type tweppy.api.API from problem 2.1
    _id: twitter id of a user.
    
    Returns
    -------

    screen name of the twitter id of type string.
    '''
    # YOUR CODE HERE
    user = _api.get_user(_id)
    
    
    return user.screen_name
get_screen_name(api, '25073877')
```

Output:
```
'realDonaldTrump'
```

<!-- grade_id: p22_tests, points: 10, grade: true, solution: false -->
```python
# https://twitter.com/realDonaldTrump
assert_true(get_screen_name(api, '25073877'), 'realDonaldTrump')

#https://twitter.com/adamdangelo
assert_true(get_screen_name(api, '10476462'), 'adamdangelo')

#https://twitter.com/BarackObama
assert_true(get_screen_name(api, '813286'), 'BarackObama')

# https://twitter.com/kc74292876
assert_true(get_screen_name(api, '839362741773156352'), 'KC74292876')
```

### Problem 2.3
Complete the function *count_tweets* by using tweepy's api wrapper to count the number of tweets for a screen name. *HINT*: you may find the function *get_screen_name* from problem 2.2 useful and you should count the number of retweets.

<!-- grade_id: p23_ans, points: 0, grade: false, solution: true -->
```python
def count_tweets(_api, _id):
    '''
    Parameters
    ----------
    _api: A Twitter API wrapper of type tweepy.api.API from problem 2.1
    _id: twitter id of a user.
    
    Returns
    -------

    number of tweets of type int
    '''
    # YOUR CODE HERE
    user = _api.get_user(get_screen_name(_api, _id))
    
    
    return user.statuses_count

```

<!-- grade_id: p23_tests, points: 10, grade: true, solution: false -->
```python
try:
    count_tweets(None, None)
    assert_true(False, "You function should make use of tweepy's api and the twitter id")
except:
    pass

count = count_tweets(api, '839362741773156352')
assert_equal(count, 19)
```

```python

```