# Week 11 Problem 3

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do not write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select *Kernel*, and restart the kernel and run all cells (*Restart & Run all*).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select *File* → *Save and CheckPoint*)

5. When you are ready to submit your assignment, go to *Dashboard* → *Assignments* and click the *Submit* button. Your work is not submitted until you click *Submit*.

6. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

## Author: Radhir Kothuri
### Primary Reviewer: Apurv Garg

# Due Date: 6 PM, April 9, 2018

```python
# testing tools
from nose.tools import (
    assert_equal, assert_true
)
from IPython.display import display_html
from bs4 import BeautifulSoup
```

```python
html = '''
<div class="container">
<h2>Countries & Capitals</h2>
<table class="two-column td-red">
<thead><tr><th>Country</th><th>Capital city</th></tr></thead><tbody>
<tr class="grey"><td>Afghanistan</td><td>Kabul</td></tr>
<tr><td>Albania</td><td>Tirana</td></tr>
<tr class="grey"><td>Algeria</td><td>Algiers</td></tr>
<tr><td>Andorra</td><td>Andorra la Vella</td></tr>
<tr class="grey"><td>Angola</td><td>Luanda</td></tr>
<tr><td>Antigua and Barbuda</td><td>Saint John's</td></tr>
<tr class="grey"><td>Argentina</td><td>Buenos Aires</td></tr>
<tr><td>Armenia</td><td>Yerevan</td></tr>
<tr class="grey"><td>Australia</td><td>Canberra</td></tr>
<tr><td>Austria</td><td>Vienna</td></tr>
<tr class="grey"><td>Azerbaijan</td><td>Baku</td></tr>
<tr class="grey"><td>Bahamas</td><td>Nassau</td></tr>
<tr><td>Bahrain</td><td>Manama</td></tr>
<tr class="grey"><td>Bangladesh</td><td>Dhaka</td></tr>
<tr><td>Barbados</td><td>Bridgetown</td></tr>
<tr class="grey"><td>Belarus</td><td>Minsk</td></tr>
<tr><td>Belgium</td><td>Brussels</td></tr>
<tr class="grey"><td>Belize</td><td>Belmopan</td></tr>
<tr><td>Benin</td><td>Porto-Novo</td></tr>
<tr class="grey"><td>Bhutan</td><td>Thimphu</td></tr>
<tr><td>Bolivia</td><td>Sucre <span>(de jure),</span> <br>La Paz <span>(seat of government)</span></td></tr>
<tr class="grey"><td>Bosnia and Herzegovina</td><td>Sarajevo</td></tr>
<tr><td>Botswana</td><td>Gaborone</td></tr>
<tr class="grey"><td>Brazil</td><td>Brasilia</td></tr>
<tr><td>Brunei</td><td>Bandar Seri Begawan</td></tr>
<tr class="grey"><td>Bulgaria</td><td>Sofia</td></tr>
<tr><td>Burkina Faso</td><td>Ouagadougou</td></tr>
<tr class="grey"><td>Burundi</td><td>Bujumbura</td></tr>
<tr class="grey"><td>Cabo Verde</td><td>Praia</td></tr>
<tr><td>Cambodia</td><td>Phnom Penh</td></tr>
<tr class="grey"><td>Cameroon</td><td>Yaounde</td></tr>
<tr><td>Canada</td><td>Ottawa</td></tr>
<tr class="grey"><td>Central African Republic</td><td>Bangui</td></tr>
<tr><td>Chad</td><td>N'Djamena</td></tr>
<tr class="grey"><td>Chile</td><td>Santiago</td></tr>
<tr><td>China</td><td>Beijing</td></tr>
<tr class="grey"><td>Colombia</td><td>Bogotá</td></tr>
<tr><td>Comoros</td><td>Moroni</td></tr>
<tr class="grey"><td><span class="black">Democratic Republic of the</span> Congo</td><td>Kinshasa</td></tr>
<tr><td><span class="black">Republic of the</span> Congo</td><td>Brazzaville</td></tr>
<tr class="grey"><td>Costa Rica</td><td>San Jose</td></tr>
<tr><td>Cote d'Ivoire</td><td>Yamoussoukro</td></tr>
<tr class="grey"><td>Croatia</td><td>Zagreb</td></tr>
<tr><td>Cuba</td><td>Havana</td></tr>
<tr class="grey"><td>Cyprus</td><td>Nicosia</td></tr>
<tr><td>Czech Republic</td><td>Prague</td></tr>
<tr class="grey"><td>Denmark</td><td>Copenhagen</td></tr>
<tr><td>Djibouti</td><td>Djibouti (city)</td></tr>
<tr class="grey"><td>Dominica</td><td>Roseau</td></tr>
<tr><td>Dominican Republic</td><td>Santo Domingo</td></tr>
<tr class="grey"><td>Ecuador</td><td>Quito</td></tr>
<tr><td>Egypt</td><td>Cairo</td></tr>
<tr class="grey"><td>El Salvador</td><td>San Salvador</td></tr>
<tr><td>Equatorial Guinea</td><td>Malabo <span>(de jure),</span> <br>Oyala <span>(seat of government)</span></td></tr>
<tr class="grey"><td>Eritrea</td><td>Asmara</td></tr>
<tr><td>Estonia</td><td>Tallinn</td></tr>
<tr class="grey"><td>Ethiopia</td><td>Addis Ababa</td></tr>
<tr class="grey"><td>Fiji</td><td>Suva</td></tr>
<tr><td>Finland</td><td>Helsinki</td></tr>
<tr class="grey"><td>France</td><td>Paris</td></tr>
<tr class="grey"><td>Gabon</td><td>Libreville</td></tr>
<tr><td>Gambia</td><td>Banjul</td></tr>
<tr class="grey"><td>Georgia</td><td>Tbilisi</td></tr>
<tr><td>Germany</td><td>Berlin</td></tr>
<tr class="grey"><td>Ghana</td><td>Accra</td></tr>
<tr><td>Greece</td><td>Athens</td></tr>
<tr class="grey"><td>Grenada</td><td>Saint George's</td></tr>
<tr><td>Guatemala</td><td>Guatemala City</td></tr>
<tr class="grey"><td>Guinea</td><td>Conakry</td></tr>
<tr><td>Guinea-Bissau</td><td>Bissau</td></tr>
<tr class="grey"><td>Guyana</td><td>Georgetown</td></tr>
<tr class="grey"><td>Haiti</td><td>Port-au-Prince</td></tr>
<tr><td>Honduras</td><td>Tegucigalpa</td></tr>
<tr class="grey"><td>Hungary</td><td>Budapest</td></tr>
<tr class="grey"><td>Iceland</td><td>Reykjavik</td></tr>
<tr><td>India</td><td>New Delhi</td></tr>
<tr class="grey"><td>Indonesia</td><td>Jakarta</td></tr>
<tr><td>Iran</td><td>Tehran</td></tr>
<tr class="grey"><td>Iraq</td><td>Baghdad</td></tr>
<tr><td>Ireland</td><td>Dublin</td></tr>
<tr class="grey"><td>Israel</td><td>Jerusalem</td></tr>
<tr><td>Italy</td><td>Rome</td></tr>
<tr class="grey"><td>Jamaica</td><td>Kingston</td></tr>
<tr><td>Japan</td><td>Tokyo</td></tr>
<tr class="grey"><td>Jordan</td><td>Amman</td></tr>
<tr class="grey"><td>Kazakhstan</td><td>Astana</td></tr>
<tr><td>Kenya</td><td>Nairobi</td></tr>
<tr class="grey"><td>Kiribati</td><td>Tarawa</td></tr>
<tr><td>Kosovo</td><td>Pristina</td></tr>
<tr class="grey"><td>Kuwait</td><td>Kuwait City</td></tr>
<tr><td>Kyrgyzstan</td><td>Bishkek</td></tr>
<tr class="grey"><td>Laos</td><td>Vientiane</td></tr>
<tr><td>Latvia</td><td>Riga</td></tr>
<tr class="grey"><td>Lebanon</td><td>Beirut</td></tr>
<tr><td>Lesotho</td><td>Maseru</td></tr>
<tr class="grey"><td>Liberia</td><td>Monrovia</td></tr>
<tr><td>Libya</td><td>Tripoli</td></tr>
<tr class="grey"><td>Liechtenstein</td><td>Vaduz</td></tr>
<tr><td>Lithuania</td><td>Vilnius</td></tr>
<tr class="grey"><td>Luxembourg</td><td>Luxembourg (city)</td></tr>
</tbody>
</table>
<table class="two-column td-red">
<tr class="grey"><td>Macedonia (FYROM)</td><td>Skopje</td></tr>
<tr><td>Madagascar</td><td>Antananarivo</td></tr>
<tr class="grey"><td>Malawi</td><td>Lilongwe</td></tr>
<tr><td>Malaysia</td><td>Kuala Lumpur</td></tr>
<tr class="grey"><td>Maldives</td><td>Male</td></tr>
<tr><td>Mali</td><td>Bamako</td></tr>
<tr class="grey"><td>Malta</td><td>Valletta</td></tr>
<tr><td>Marshall Islands</td><td>Majuro</td></tr>
<tr class="grey"><td>Mauritania</td><td>Nouakchott</td></tr>
<tr><td>Mauritius</td><td>Port Louis</td></tr>
<tr class="grey"><td>Mexico</td><td>Mexico City</td></tr>
<tr><td>Micronesia</td><td>Palikir</td></tr>
<tr class="grey"><td>Moldova</td><td>Chisinau</td></tr>
<tr><td>Monaco</td><td>Monaco</td></tr>
<tr class="grey"><td>Mongolia</td><td>Ulaanbaatar</td></tr>
<tr><td>Montenegro</td><td>Podgorica</td></tr>
<tr class="grey"><td>Morocco</td><td>Rabat</td></tr>
<tr><td>Mozambique</td><td>Maputo</td></tr>
<tr class="grey"><td>Myanmar (Burma)</td><td>Naypyidaw</td></tr>
<tr class="grey"><td>Namibia</td><td>Windhoek</td></tr>
<tr><td>Nauru</td><td>Yaren District <span>(de facto)</span></td></tr>
<tr class="grey"><td>Nepal</td><td>Kathmandu</td></tr>
<tr><td>Netherlands</td><td>Amsterdam</td></tr>
<tr class="grey"><td>New Zealand</td><td>Wellington</td></tr>
<tr><td>Nicaragua</td><td>Managua</td></tr>
<tr class="grey"><td>Niger</td><td>Niamey</td></tr>
<tr><td>Nigeria</td><td>Abuja</td></tr>
<tr class="grey"><td>North Korea</td><td>Pyongyang</td></tr>
<tr><td>Norway</td><td>Oslo</td></tr>
<tr class="grey"><td>Oman</td><td>Muscat</td></tr>
<tr class="grey"><td>Pakistan</td><td>Islamabad</td></tr>
<tr><td>Palau</td><td>Ngerulmud</td></tr>
<tr class="grey"><td>Palestine</td><td>Jerusalem (East)</td></tr>
<tr><td>Panama</td><td>Panama City</td></tr>
<tr class="grey"><td>Papua New Guinea</td><td>Port Moresby</td></tr>
<tr><td>Paraguay</td><td>Asunción</td></tr>
<tr class="grey"><td>Peru</td><td>Lima</td></tr>
<tr><td>Philippines</td><td>Manila</td></tr>
<tr class="grey"><td>Poland</td><td>Warsaw</td></tr>
<tr><td>Portugal</td><td>Lisbon</td></tr>
<tr class="grey"><td>Qatar</td><td>Doha</td></tr>
<tr class="grey"><td>Romania</td><td>Bucharest</td></tr>
<tr><td>Russia</td><td>Moscow</td></tr>
<tr class="grey"><td>Rwanda</td><td>Kigali</td></tr>
<tr class="grey"><td>Saint Kitts and Nevis</td><td>Basseterre</td></tr>
<tr><td>Saint Lucia</td><td>Castries</td></tr>
<tr class="grey"><td>Saint Vincent and the Grenadines</td><td>Kingstown</td></tr>
<tr><td>Samoa</td><td>Apia</td></tr>
<tr class="grey"><td>San Marino</td><td>San Marino</td></tr>
<tr><td>Sao Tome and Principe</td><td>São Tomé</td></tr>
<tr class="grey"><td>Saudi Arabia</td><td>Riyadh</td></tr>
<tr><td>Senegal</td><td>Dakar</td></tr>
<tr class="grey"><td>Serbia</td><td>Belgrade</td></tr>
<tr><td>Seychelles</td><td>Victoria</td></tr>
<tr class="grey"><td>Sierra Leone</td><td>Freetown</td></tr>
<tr><td>Singapore</td><td>Singapore</td></tr>
<tr class="grey"><td>Slovakia</td><td>Bratislava</td></tr>
<tr><td>Slovenia</td><td>Ljubljana</td></tr>
<tr class="grey"><td>Solomon Islands</td><td>Honiara</td></tr>
<tr><td>Somalia</td><td>Mogadishu</td></tr>
<tr class="grey"><td>South Africa</td><td>Pretoria <span>(administrative),</span> <br>Cape Town <span>(legislative),</span> <br>Bloemfontein <span>(judicial)</span></td></tr>
<tr><td>South Korea</td><td>Seoul</td></tr>
<tr class="grey"><td>South Sudan</td><td>Juba</td></tr>
<tr><td>Spain</td><td>Madrid</td></tr>
<tr class="grey"><td>Sri Lanka</td><td>Sri Jayawardenepura Kotte</td></tr>
<tr><td>Sudan</td><td>Khartoum</td></tr>
<tr class="grey"><td>Suriname</td><td>Paramaribo</td></tr>
<tr><td>Swaziland</td><td>Mbabane <span>(administrative),</span> <br>Lobamba <span>(legislative, royal)</span></td></tr>
<tr class="grey"><td>Sweden</td><td>Stockholm</td></tr>
<tr><td>Switzerland</td><td>Bern</td></tr>
<tr class="grey"><td>Syria</td><td>Damascus</td></tr>
<tr class="grey"><td>Taiwan</td><td>Taipei</td></tr>
<tr><td>Tajikistan</td><td>Dushanbe</td></tr>
<tr class="grey"><td>Tanzania</td><td>Dodoma</td></tr>
<tr><td>Thailand</td><td>Bangkok</td></tr>
<tr class="grey"><td>Timor-Leste</td><td>Dili</td></tr>
<tr><td>Togo</td><td>Lomé</td></tr>
<tr class="grey"><td>Tonga</td><td>Nukuʻalofa</td></tr>
<tr><td>Trinidad and Tobago</td><td>Port of Spain</td></tr>
<tr class="grey"><td>Tunisia</td><td>Tunis</td></tr>
<tr><td>Turkey</td><td>Ankara</td></tr>
<tr class="grey"><td>Turkmenistan</td><td>Ashgabat</td></tr>
<tr><td>Tuvalu</td><td>Funafuti</td></tr>
<tr class="grey"><td>Uganda</td><td>Kampala</td></tr>
<tr><td>Ukraine</td><td>Kyiv <span>(also known as Kiev)</span></td></tr>
<tr class="grey"><td>United Arab Emirates</td><td>Abu Dhabi</td></tr>
<tr><td>United Kingdom</td><td>London</td></tr>
<tr class="grey"><td>United States of America</td><td>Washington, D.C.</td></tr>
<tr><td>Uruguay</td><td>Montevideo</td></tr>
<tr class="grey"><td>Uzbekistan</td><td>Tashkent</td></tr>
<tr class="grey"><td>Vanuatu</td><td>Port Vila</td></tr>
<tr><td>Vatican City (Holy See)</td><td>Vatican City</td></tr>
<tr class="grey"><td>Venezuela</td><td>Caracas</td></tr>
<tr><td>Vietnam</td><td>Hanoi</td></tr>
<tr class="grey"><td>Yemen</td><td>Sana'a</td></tr>
<tr class="grey"><td>Zambia</td><td>Lusaka</td></tr>
<tr><td>Zimbabwe</td><td>Harare</td></tr>
</tbody>
</table>
</div>
'''
display_html(html, raw=True)
with open('countries.html', 'w') as fout:
    fout.write(html)
```

## Question 1:  Parsing HTML file using Beautiful Soup

The above cell contains html text written to `countries.html` that contains a table that with all the countries in the world and their respective capital cities. Using Beautiful Soup, complete the function `countries_map()` and return a dictionary that maps each capital to its respective country. For example, by observing the html above, it is evident that the return dictionary should be formatted as such: `{'Kabul':'Afghanistan', 'Tirana': 'Albania', 'Algiers': 'Algeria', ...}`.

<!-- grade_id: problem1_answer, points: 0, grade: false, solution: true -->
```python
def countries_map(html_file):
    """
    Construct a dictionary mapping the country to capital 
    in the above countries.html file
    
    Parameters
    ----------
    html_file: a string that represents the filepath to the html file to be parsed
    
    Returns
    -------
    A dictionary
    """
    
    # YOUR CODE HERE

    html = open(html_file, 'r')
    soup = BeautifulSoup(html, 'lxml')
    tables = soup.find_all('table')
    
    country = []
    captail = []
    
    countries = {}
    
    i = 0
    for table in tables:

        for tdx in table.find_all('td'):
            if i % 2 == 0:
                country.append(tdx.contents[0])
            else:
                captail.append(tdx.contents[0])
            i += 1

    countries = dict(zip(captail, country))
    html.close()
    return countries

```

<!-- grade_id: problem1_test, points: 10, grade: true, solution: false -->
```python
countries = countries_map('countries.html')
assert_equal(countries['Kabul'], 'Afghanistan')
assert_equal(countries['Washington, D.C.'], 'United States of America')
assert_equal(countries['Buenos Aires'], 'Argentina')
assert_equal(len(countries), 197)
```

```python
# This code is for problem 2 and 3
# It makes a request to a github repository that hosts an html file and then renders the HTML 
import urllib.request
with urllib.request.urlopen("https://raw.githubusercontent.com/kothuri2/info490-w5p1/master/pokedex.html") as url:
    html = url.read()
display_html(html, raw=True)
```

## Question 2

In this question, we will be implementing a simple parsing function, that will return to us the most powerful pokemon defined **greater than or equal to above a certain threshold**. We will quantify the power of a pokemon, in terms of the `total` attribute for each pokemon (i.e. the 4th column value). Finish the function below, `get_pokemon_above_threshold` that returns a list of pokemon names (list of strings) and takes in 2 function parametesr: `threshold`, an integer and `html_file`, a string that represents the file path of the html file to be parsed.

- For example, given the first 5 pokemon and if the function call was `get_pokemon_above_threshold(350)`, then `[Ivysaur, Venusaur, Charmeleon]` would be returned since these are the 3 pokemon with totals **greater than or equal to** 350.
- Note that `Venusaur` is only added once, even though there is a second version of Venusaur with a higher total. Therefore, do not double count the same pokemon with a `Mega` prefix.

<!-- grade_id: problem2_answer, points: 0, grade: false, solution: true -->
```python
def get_pokemon_above_threshold(threshold, html_file):
    '''    
    Return a list of pokemon names above the threshold
    
    Parameters
    ----------
    threshold: an int
    html_file: a string that represents the filepath to the html file to be parsed
    
    Returns
    -------
    A list of strings
    '''
    
    # YOUR CODE HERE
    html = open(html_file)
    soup = BeautifulSoup(html, 'lxml')
    tb = soup.find('table').tbody

    total = []
    pokemon = []
    
    for row in tb.find_all('tr'):   
        tdx = [val for val in row.find_all('td')]
        if (tdx[1].a.contents[0] not in pokemon) & \
        (int(tdx[3].contents[0].strip()) >= threshold):
            pokemon.append(tdx[1].a.contents[0].strip())
            
#             total.append(tdx[3].contents[0].strip()
            
    return pokemon

    
```

Output:
```
['Bulbasaur',
 'Ivysaur',
 'Venusaur',
 'Charmeleon',
 'Charizard',
 'Squirtle',
 'Wartortle',
 'Blastoise',
 'Butterfree',
 'Beedrill',
 'Pidgeotto',
 'Pidgeot',
 'Raticate',
 'Fearow',
 'Arbok',
 'Pikachu',
 'Raichu',
 'Sandslash',
 'Nidorina',
 'Nidoqueen',
 'Nidorino',
 'Nidoking',
 'Clefairy',
 'Clefable',
 'Ninetales',
 'Wigglytuff',
 'Golbat',
 'Oddish',
 'Gloom',
 'Vileplume',
 'Parasect',
 'Venomoth',
 'Dugtrio',
 'Persian',
 'Psyduck',
 'Golduck',
 'Primeape',
 'Growlithe',
 'Arcanine',
 'Poliwhirl',
 'Poliwrath',
 'Abra',
 'Kadabra',
 'Alakazam',
 'Machoke',
 'Machamp',
 'Weepinbell',
 'Victreebel',
 'Tentacool',
 'Tentacruel',
 'Graveler',
 'Golem',
 'Ponyta',
 'Rapidash',
 'Slowpoke',
 'Slowbro',
 'Magnemite',
 'Magneton',
 "Farfetch'd",
 'Doduo',
 'Dodrio',
 'Seel',
 'Dewgong',
 'Grimer',
 'Muk',
 'Cloyster',
 'Gastly',
 'Haunter',
 'Gengar',
 'Onix',
 'Drowzee',
 'Hypno',
 'Krabby',
 'Kingler',
 'Voltorb',
 'Electrode',
 'Exeggcute',
 'Exeggutor',
 'Cubone',
 'Marowak',

```

<!-- grade_id: problem2_test, points: 10, grade: true, solution: false -->
```python
filename = '/home/data_scientist/data/misc/pokedex.html'
pokemon_thresholds = get_pokemon_above_threshold(310, filename)
assert_true('Bulbasaur' in pokemon_thresholds)
assert_true('Venusaur' in pokemon_thresholds)
assert_true('Charmander' not in pokemon_thresholds)
assert_true(len(pokemon_thresholds) == 639)

pokemon_thresholds = get_pokemon_above_threshold(700, filename)
assert_true('Tyranitar' in pokemon_thresholds)
assert_true('Salamence' in pokemon_thresholds)
assert_true('Metagross' in pokemon_thresholds)
assert_true('Latios' in pokemon_thresholds)
assert_true('Latias' in pokemon_thresholds)
assert_true('Charmander' not in pokemon_thresholds)
assert_true(len(pokemon_thresholds) == 14)
```

## Question 3

In the question, we will be doing something very similar to the previous question. We will be returning a list of pokemon (list of strings) of a particular `type`. Complete the function `get_pokemon_by_type` that takes in 2 function parameters `p_type` (a string) and returns a list of pokemon names that match `type`, and `html_file`, a string that represents the file path of the html file to be parsed.

- For example, given the first 5 pokemon and if the function call was `get_pokemon_by_type(`Grass`)`, then `[Bulbasaur, Ivysaur, Venusaur]` would be returned since these are the 3 pokemon that are of `Grass` type.
- Note that `Venusaur` is only added once, even though there is a second version of Venusaur of `Grass` type. Theregore, do not double count the same pokemon with a `Mega` prefix.

<!-- grade_id: problem3_answer, points: 0, grade: false, solution: true -->
```python
def get_pokemon_by_type(p_type, html_file):
    '''    
    Return a list of pokemon names that match `p_type`
    
    Parameters
    ----------
    p_type: a string
    html_file: a string that represents the filepath to the html file to be parsed
    
    Returns
    -------
    A list of strings
    '''
    
    # YOUR CODE HERE
    html = open(html_file)
    soup = BeautifulSoup(html, 'lxml')
    tb = soup.find('table').tbody

    total = []
    pokemon = []
    
    for row in tb.find_all('tr'):
        tdx = [val for val in row.find_all('td')]
        for i in tdx[2].find_all('a'):
            if (tdx[1].a.contents[0] not in pokemon) & (i.contents[0] == p_type):
                pokemon.append(tdx[1].a.contents[0].strip())
#         print(tdx[2].find_all('a')[1].contents)
#         if (tdx[1].a.contents[0] not in pokemon) & \
#         (tdx[2].a.contents[0].strip() == p_type):
#             pokemon.append(tdx[1].a.contents[0].strip())
            
#             total.append(tdx[3].contents[0].strip()
            
    return pokemon

```

Output:
```
['Bulbasaur',
 'Ivysaur',
 'Venusaur',
 'Oddish',
 'Gloom',
 'Vileplume',
 'Paras',
 'Parasect',
 'Bellsprout',
 'Weepinbell',
 'Victreebel',
 'Exeggcute',
 'Exeggutor',
 'Tangela',
 'Chikorita',
 'Bayleef',
 'Meganium',
 'Bellossom',
 'Hoppip',
 'Skiploom',
 'Jumpluff',
 'Sunkern',
 'Sunflora',
 'Celebi',
 'Treecko',
 'Grovyle',
 'Sceptile',
 'Lotad',
 'Lombre',
 'Ludicolo',
 'Seedot',
 'Nuzleaf',
 'Shiftry',
 'Shroomish',
 'Breloom',
 'Roselia',
 'Cacnea',
 'Cacturne',
 'Lileep',
 'Cradily',
 'Tropius',
 'Turtwig',
 'Grotle',
 'Torterra',
 'Budew',
 'Roserade',
 'Wormadam',
 'Cherubi',
 'Cherrim',
 'Carnivine',
 'Snover',
 'Abomasnow',
 'Tangrowth',
 'Leafeon',
 'Rotom',
 'Shaymin',
 'Snivy',
 'Servine',
 'Serperior',
 'Pansage',
 'Simisage',
 'Sewaddle',
 'Swadloon',
 'Leavanny',
 'Cottonee',
 'Whimsicott',
 'Petilil',
 'Lilligant',
 'Maractus',
 'Deerling',
 'Sawsbuck',
 'Foongus',
 'Amoonguss',
 'Ferroseed',
 'Ferrothorn',
 'Virizion',
 'Chespin',
 'Quilladin',
 'Chesnaught',
 'Sk
```

<!-- grade_id: problem3_test, points: 10, grade: true, solution: false -->
```python
filename = '/home/data_scientist/data/misc/pokedex.html'
pokemon_types = get_pokemon_by_type('Grass', filename)
assert_true('Bulbasaur' in pokemon_types)
assert_true('Venusaur' in pokemon_types)
assert_true('Charmander' not in pokemon_types)
assert_true(len(pokemon_types) == 98)

pokemon_types = get_pokemon_by_type('Water', filename)
assert_true('Wartortle' in pokemon_types)
assert_true('Blastoise' in pokemon_types)
assert_true('Squirtle' in pokemon_types)
assert_true('Psyduck' in pokemon_types)
assert_true('Golduck' in pokemon_types)
assert_true('Charmander' not in pokemon_types)
assert_true(len(pokemon_types) == 132)
```

```python

```