# Week 13 Problem 3

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says YOUR CODE HERE. Do not write your answer in anywhere else other than where it says YOUR CODE HERE. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select Kernel, and restart the kernel and run all cells (Restart & Run all).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select File → Save and CheckPoint)

5. When you are ready to submit your assignment, go to Dashboard → Assignments and click the Submit button. Your work is not submitted until you click Submit.

6. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

7. If your code does not pass the unit tests, it will not pass the autograder.


## Author: Apurv Garg
### Primary Reviewer: Radhir Kothuri


# Due Date: 6 PM, April 23, 2018

```python
# Set up Notebook

% matplotlib inline

# Standard imports
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
import seaborn as sns
from numpy.random import RandomState
from nose.tools import assert_false, assert_equal, assert_almost_equal, assert_true, assert_in, assert_is_not

# These two lines suppress warnings that sometimes
# occur when making visualizations
import warnings
warnings.filterwarnings('ignore')
```

# Problem 1

For this problem, complete the function `integral1` which will take `pts` as parameter and return the integral value and the function of interest that is created for purpose of integration. 

Calculate the value of following using MCMC: $$\int_{0}^{1} | xln(x)e^{x} | dx$$

- Use RandomState(23) to generate random uniform numbers and number of points that should be generated should be equal to `pts` parameter.
- Create Boolean mask array from function of interest and compute the integral(similar to how we did in the notebook)


<!-- grade_id: problem1_answer, points: 0, grade: false, solution: true -->
```python
def integral1(pts):
    '''           
    Parameters
    ----------
    pts : Number of sample points
    
    Returns
    -------
    A tuple of 2 containing integral value, 
    function of interest created for purpose of integration
    '''
    # YOUR CODE HERE
    rng = RandomState(23)
    x, y = rng.uniform(0,1,size=(2,pts))
    
    def the_func(x):
        return np.abs(x * np.log(x) * np.exp(x))
    mask = y < the_func(x)
    value = y[mask].shape[0] / pts
    
    return value, the_func
```

<!-- grade_id: problem1_tests, points: 5, grade: true, solution: false -->
```python
val1, function1= integral1(100000)
assert_almost_equal(val1, 0.4008, 4)
val2, function1= integral1(100)
assert_equal(val2, 0.42)
assert_almost_equal(function1(2), 10.243406803946097, 4)
assert_almost_equal(function1(3), 66.198653064933083, 4)
```

# Methodology for below Problems

You can use the following approach to calculate the below integrals:

- Generate uniform numbers from [a,b] (that is lower integral limit to upper integral limit) of size `pts`.
- y should be the value returned by passing value of x to the function of interest that is created for purpose of integration.
- Integral value will be equal to: $$\frac{(\sum y)*(b-a)}{\#pts}$$	where b is upperlimit and a is lowerlimit of integral.

# Problem 2

For this problem, complete the function `integral2` which will take `pts` as parameter and return the integral value and the function of interest that is created for purpose of integration. 

Calculate the value of following using MCMC: $$\int_{0}^{2\pi} sin(x)e^{-x} dx$$ 

- Use RandomState(23) to generate random uniform numbers and number of points that should be generated should be equal to `pts` parameter.

<!-- grade_id: problem2_answer, points: 0, grade: false, solution: true -->
```python
def integral2(pts):
    '''           
    Parameters
    ----------
    pts : Number of sample points
    
    Returns
    -------
    A tuple of 2 containing integral value, 
    function of interest created for purpose of integration
    '''
    # YOUR CODE HERE
    rng = RandomState(23)
    x, y = rng.uniform(0,2*np.pi,size=(2,pts))
    
    def the_func(x):
        return np.sin(x) * np.exp(-x)
#     mask = y < the_func(x)
    y1 = the_func(x)
#     value = y[mask].shape[0] * 2*np.pi/ pts
    value = sum(y1) * 2 * np.pi / pts
    return value, the_func

```

<!-- grade_id: problem2_tests, points: 5, grade: true, solution: false -->
```python
val1_p2, function2= integral2(100000)
assert_almost_equal(val1_p2, 0.50085896, 4)
val2_p2, function2= integral2(1000000)
assert_almost_equal(val2_p2, 0.4991577604, 4)
val3_p2, function2= integral2(100)
assert_almost_equal(val3_p2, 0.5533328458, 4)
assert_almost_equal(function2(2), 0.12306002480577674, 4)
assert_almost_equal(function2(1), 0.30955987565311222, 4)
```

# Problem 3

For this problem, complete the function `integral3` which will take `pts` as parameter and return the integral value and the function of interest that is created for purpose of integration. 

- Calculate the value of following using MCMC:  $$\int_{0}^{\infty} \frac{3}{4}x^4e^{-x^{\frac{3}{4}}} dx$$  

- Use RandomState(23) to generate random uniform numbers and number of points that should be generated should be equal to `pts` parameter.

- Use upper limit to be 50 instead of $\infty$. This is shown in Approach1 of this [answer](https://math.stackexchange.com/a/1200465).

- The given integral is actually the value of $\Gamma(\frac{20}{3}) = 389.03492$ where $\Gamma(a) = \int_{0}^{\infty} x^{a-1}e^{-x}dx$. $\Gamma$ represents the [Gamma function](https://en.wikipedia.org/wiki/Gamma_function).

```python
3^3
```

Output:
```
0
```

<!-- grade_id: problem3_answer, points: 0, grade: false, solution: true -->
```python
def integral3(pts):
    '''           
    Parameters
    ----------
    pts : Number of sample points
    
    Returns
    -------
    A tuple of 2 containing integral value, 
    function of interest created for purpose of integration
    '''
    # YOUR CODE HERE
    rng = RandomState(23)
    x, y = rng.uniform(0, 50, size=(2,pts))
    
    def the_func(x): 
        n = 3/4 * x**4 * np.exp(-x**(3/4))
        return n
#     mask = y < the_func(x)
    y1 = the_func(x)
#     value = y[mask].shape[0] * 2*np.pi/ pts
    value = sum(y1) * 50 / pts
    return value, the_func
```

<!-- grade_id: problem3_tests, points: 10, grade: true, solution: false -->
```python
val1_p3, function3= integral3(1000000)
assert_almost_equal(val1_p3, 389.12877550513804, 4)
val2_p3, function3= integral3(100000)
assert_almost_equal(val2_p3, 391.18915255864755, 4)
assert_almost_equal(function3(3), 6.216829345559697, 4)
assert_almost_equal(function3(2), 2.2324816612309828, 4)
```

# Problem 4

For this problem, complete the function `integral4` which will take `pts` as parameter and return the integral value and the function of interest that is created for purpose of integration. 

- Calculate the value of following using MCMC: $$\int_{0}^{\infty} \frac{3}{4}x^4e^{-x^{\frac{3}{4}}} dx$$ 

using the substitution: $$u = \frac{1}{1+x}$$

- Use RandomState(23) to generate random uniform numbers and number of points that should be generated should be equal to `pts` parameter.

- The function of interest should be the one obtained after the above substitution i.e. we will be integrating with respect to $du$: $$\int_{0}^{\infty} \frac{3}{4}x^4e^{-x^{\frac{3}{4}}} dx = \int_{0}^{1} \frac{3(1-u)^4e^{-(\frac{1-u}{u})^{\frac{3}{4}}}}{4u^6} du$$


- We we be using $\int_{0}^{1} \frac{3(1-u)^4e^{-(\frac{1-u}{u})^{\frac{3}{4}}}}{4u^6}du$ to calculate the value of $\Gamma(\frac{20}{3})$ for Problem 4.

- This is shown in Approach3 of this [answer](https://math.stackexchange.com/a/1200465).


<!-- grade_id: problem4_answer, points: 0, grade: false, solution: true -->
```python
def integral4(pts):

    '''           
    Parameters
    ----------
    pts : Number of sample points
    
    Returns
    -------
    A tuple of 2 containing integral value, 
    function of interest created for purpose of integration
    '''
    # YOUR CODE HERE
    rng = RandomState(23)
    x, y = rng.uniform(0, 1, size=(2,pts))
    
    def the_func(x): 
        n = (3 * ((1-x)**4) * np.exp(-((1-x)/x)**(3/4)) )/(4 * (x ** 6))
        return n
#     mask = y < the_func(x)
    y1 = the_func(x)
#     value = y[mask].shape[0] * 2*np.pi/ pts
    value = sum(y1) * 1 / pts
    return value, the_func
```

Output:
```
(388.710212046084, <function __main__.integral4.<locals>.the_func>)
```

<!-- grade_id: problem4_tests, points: 5, grade: true, solution: false -->
```python
val1_p4, function4= integral4(1000000)
assert_almost_equal(val1_p4, 388.71021204610099, 4)
val2_p4, function4= integral4(100000)
assert_almost_equal(val2_p4, 388.23741983842672, 4)
assert_almost_equal(np.absolute(function4(3)), 0.027734664499464268, 4)
assert_almost_equal(np.absolute(function4(5)), 0.022348577866342166, 4)
```

# Problem 5

For this problem, complete the function `integral5` which will take `pts` as parameter and return the integral value and the function of interest that is created for purpose of integration. 

Calculate the value of following using MCMC: $$\int_{10}^{50} \frac{3}{4}x^4e^{-x^{\frac{3}{4}}} dx$$ 

- Use RandomState(23) to generate random uniform numbers and number of points that should be generated should be equal to `pts` parameter.

<!-- grade_id: problem5_answer, points: 0, grade: false, solution: true -->
```python
def integral5(pts):
    '''           
    Parameters
    ----------
    pts : Number of sample points
    
    Returns
    -------
    A tuple of 2 containing integral value, 
    function of interest created for purpose of integration
    '''
    # YOUR CODE HERE
    rng = RandomState(23)
    x, y = rng.uniform(10, 50, size=(2,pts))
    
    def the_func(x): 
        n = 3/4 * (x**4) * np.exp(-x**(3/4)) 
        return n
#     mask = y < the_func(x)
    y1 = the_func(x)
#     value = y[mask].shape[0] * 2*np.pi/ pts
    value = sum(y1) * 40 / pts
    return value, the_func
   
```

<!-- grade_id: problem5_tests, points: 5, grade: true, solution: false -->
```python
val1_p5, function5= integral5(1000000)
assert_almost_equal(val1_p5, 239.23033786913984, 4)
val2_p5, function5= integral5(100000)
assert_almost_equal(val2_p5, 239.49732273577268, 4)
assert_almost_equal(function5(2), 2.2324816612309828, 4)
assert_almost_equal(function5(3), 6.216829345559697, 4)
```

```python

```