# Week 13 Problem 2

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do not write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select *Kernel*, and restart the kernel and run all cells (*Restart & Run all*).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select *File* → *Save and CheckPoint*)

5. When you are ready to submit your assignment, go to *Dashboard* → *Assignments* and click the *Submit* button. Your work is not submitted until you click *Submit*.

6. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

## Author: Radhir Kothuri
### Primary Reviewer: Apurv Garg

# Due Date: 6 PM, April 23, 2018

```python
# Standard imports
import numpy as np
from numpy.random import RandomState
# testing tools
from nose.tools import (
    assert_equal, assert_true
)

from numpy.testing import assert_almost_equal

# These two lines suppress warnings that sometimes
# occur when making visualizations
import warnings
warnings.filterwarnings('ignore')
```

## Question 1

In this question, we will use discrete markov chains in order to calculate all the configurations up to an inputted number of `num_steps`. In effect, you will be implementing the discrete markov chain algorithm with inputs: `initial_state`, `state_transitions`, and determining all the intermediate configurations from 0 to `num_days`, with 0 being the `initial_state`.

- Complete the function `discrete_state` that takes in 3 function parameters: `num_steps`, `initial_state`, and `state_transitions`
- Return a 2D list of all the configurations (state probabilties) from 0 to `num_steps`. The length of the list returned should be of size `num_steps`
- The `state_transitions` will be a 2D array with size: 3 columns and 3 rows (i.e. the number of discrete states in the markov chain will be 3).
- The `initial_state` should be at index 0 in the list of state probabilities returned.
- The output when printed out in table format should look something like this:
![probabilities](images/data.png)

<!-- grade_id: problem1_answer, points: 0, grade: false, solution: true -->
```python
def discrete_state(num_steps, initial_state, state_transitions):
    """
    Return an array of the final state probabilities over the num_steps
    The size of the returned array should be of length num_steps
    
    Parameters
    ----------
    num_steps: an int
    initial_state: a list
    state_transitions: a 2d list
    
    Returns
    -------
    A 2d array
    """
    
    # YOUR CODE HERE
    state = np.zeros((num_steps, 3))
    state[0] = initial_state
    
    for i in range(1, num_steps):
        state[i] = state[i-1].dot(state_transitions)
        
    return state
```

<!-- grade_id: problem1_test, points: 10, grade: true, solution: false -->
```python
state_transitions = [[0.5, 0.25, 0.25],
                     [0.5, 0.0, 0.5],
                     [0.25, 0.25, 0.5]]
final_states = discrete_state(10, [0, 0, 1], state_transitions)
assert_equal(len(final_states), 10)
assert_equal(final_states[0][0], 0)
assert_equal(final_states[0][1], 0)
assert_equal(final_states[0][2], 1)
for array in final_states:
    assert_equal(len(array), 3)
```

```python
# View final state transitions
state_transitions = [[0.5, 0.25, 0.25],
                     [0.5, 0.0, 0.5],
                     [0.25, 0.25, 0.5]]
final_states = discrete_state(10, [0, 0, 1], state_transitions)
for idx in range(0, 10):
    # Display new state probabilities
    print(f'{idx:3d}  {final_states[idx][0]:5.4f}   {final_states[idx][1]:5.4f}    {final_states[idx][2]:5.4f}')
```

Output:
```
0  0.0000   0.0000    1.0000
  1  0.2500   0.2500    0.5000
  2  0.3750   0.1875    0.4375
  3  0.3906   0.2031    0.4062
  4  0.3984   0.1992    0.4023
  5  0.3994   0.2002    0.4004
  6  0.3999   0.2000    0.4001
  7  0.4000   0.2000    0.4000
  8  0.4000   0.2000    0.4000
  9  0.4000   0.2000    0.4000
```

## Question 2

In this question we will determine during which step in the iterations, does the markov chain reach its consistent final state. We define the consistent final state as the first state, `i`, in the list of state probabilties where each consecutive state probabilty is the same (i.e. states [i+1, `num_steps`] have the same state probabilities as step, `i`.

For example, given this following state transition probabilties with `num_steps = 10`, we can say that the markov chain reaches its consistent final state at state 7 as this is the first state where all the states after have the same probabilties as step 7.

![probabilities](images/data.png)

- Complete the function `determine_consistent_state` below which takes in 1 function parameter: `final_states` a 2D array of final state configurations and return an integer that represents the consistent final state as described above.
- **Note: The final states can be any of size `nxn` where n is the number of states in the markov chain.**

<!-- grade_id: problem2_answer, points: 0, grade: false, solution: true -->
```python
def determine_consistent_state(final_states):
    '''    
    Return an integer that represents the consistent final state in the final_states
    
    Parameters
    ----------
    final_states: a 2D array of final_state configurations
    
    Returns
    -------
    An integer
    '''
    
    # YOUR CODE HERE
    n = 0
    for i in range(0, len(final_states)):
        if i == len(final_states)-1:
            n = i - 1
        elif sum(final_states[i] == final_states[i + 1])==3:
            n = i
            break
        
    return n + 1
  
```

<!-- grade_id: problem2_test, points: 10, grade: true, solution: false -->
```python
state_transitions = [[0.5, 0.25, 0.25],
                     [0.5, 0.0, 0.5],
                     [0.25, 0.25, 0.5]]
final_states = np.around(discrete_state(10, [0, 0, 1], state_transitions), 3)
consistent_state = determine_consistent_state(final_states)
assert_equal(consistent_state, 7)

state_transitions = [[0.50, 0.35, 0.15], 
                     [0.30, 0.50, 0.20], 
                     [0.40, 0.20, 0.40]]
final_states = np.around(discrete_state(10, [0, 0, 1], state_transitions), 3)
consistent_state = determine_consistent_state(final_states)
assert_equal(consistent_state, 7)

consistent_state = determine_consistent_state(
               np.around(
                   [[ 0.26425401,  0.45496578,  0.58154822,  0.62245072],
                   [ 0.42417175,  0.19164469,  0.68186883,  0.01022666],
                   [ 0.17224821,  0.39245854,  0.9726721 ,  0.12367582],
                   [ 0.50411052,  0.70906992,  0.0166034 ,  0.41371041],
                   [ 0.50411067,  0.70907015,  0.0165674 ,  0.41372987]],
               3))
assert_equal(consistent_state, 4)
```

## Question 3

In this question, we will be doing something very similar to problem 1, but we will be exploring continuous markov chains instead.

- Complete the function `continuous_state` that takes in 2 function parameters: `num_steps` and `num_chains`.
- Use the continuous markov chain algorithm as described in the lesson in order to generate the first `num_steps` state transition probabilities.
- Use 18 for the `RandomState` generator
- The first 2 parameters to the `normal()` function for the `RandomState` generator correspond to the `mean` and `std` respectively. Use a factor of `0.4` for the `mean` parmater and a value of `1` for the `std`.
- Note: **For the initial state only, do not use `0.4` for the mean parameter and `1` for the `std`. Simply just use the output of the `normal()` function.**
- Return a 2D list of the first `num_steps` state probabilities.


<!-- grade_id: problem3_answer, points: 0, grade: false, solution: true -->
```python
def continuous_state(num_steps, num_chains):
    """
    Return an array of the final state probabilities over the num_steps
    The size of the returned array should be of length num_steps
    
    Parameters
    ----------
    num_steps: an int
    num_chains: an int
    
    Returns
    -------
    A 2d array
    """
    
    # YOUR CODE HERE
    rng = RandomState(18)
    
    state = np.zeros((num_steps, num_chains))
    
    state[0] = rng.normal()
    
    for i in range(1, num_steps):
        state[i] = rng.normal(0.4*state[i-1], 1, num_chains)
        
    return state
```

<!-- grade_id: problem3_test, points: 10, grade: true, solution: false -->
```python
final_states = continuous_state(10, 4)
assert_equal(len(final_states), 10)
for array in final_states:
    assert_equal(len(array), 4)
assert_almost_equal(final_states[0][0], 0.0794, decimal=4)
assert_almost_equal(final_states[0][1], 0.0794, decimal=4)
assert_almost_equal(final_states[0][2], 0.0794, decimal=4)

assert_almost_equal(final_states[1][0], 2.2220, decimal=4)
assert_almost_equal(final_states[1][1], -0.1031, decimal=4)
assert_almost_equal(final_states[1][2], 0.1923, decimal=4)

final_states = continuous_state(5000, 4)
assert_equal(len(final_states), 5000)
for array in final_states:
    assert_equal(len(array), 4)
```

```python
# View final state transitions
final_states = continuous_state(10, 4)
for idx in range(0, 10):
    # Display new state probabilities
    print(f'{idx:3d}  {final_states[idx][0]:5.4f}   {final_states[idx][1]:5.4f}    {final_states[idx][2]:5.4f}')
```

Output:
```
0  0.0794   0.0794    0.0794
  1  2.2220   -0.1031    0.1923
  2  1.5122   0.9677    0.4712
  3  -0.0410   1.1925    0.6049
  4  -1.1181   -0.5177    1.5875
  5  0.0900   1.8599    0.3447
  6  -0.5285   -0.4240    -0.8841
  7  -0.7452   -1.5200    -0.1236
  8  1.0727   1.8557    1.0785
  9  0.4965   1.8195    -0.1222
```

```python

```