# Week 2 Problem 2

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says YOUR CODE HERE. Do not write your answer in anywhere else other than where it says YOUR CODE HERE. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select Kernel, and restart the kernel and run all cells (Restart & Run all).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select File → Save and CheckPoint)

5. When you are ready to submit your assignment, go to Dashboard → Assignments and click the Submit button. Your work is not submitted until you click Submit.

6. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

7. If your code does not pass the unit tests, it will not pass the autograder.



# Due Date: 6 PM, January 29, 2018


```python
# Set up Notebook
% matplotlib inline

# Standard imports
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
import seaborn as sns
from nose.tools import assert_false, assert_equal, assert_almost_equal, assert_true
from sklearn.metrics import mean_absolute_error
from sklearn.metrics import mean_squared_error
from sklearn import metrics
from nose.tools import assert_equal
from sklearn.model_selection import train_test_split
from sklearn.tree import DecisionTreeClassifier
from sklearn.tree import DecisionTreeRegressor

# We do this to ignore several specific Pandas warnings
import warnings
warnings.filterwarnings('ignore')

sns.set_style('white')

# Set up Notebook
% matplotlib inline
```

### Reading in the Titanic Data Set


This data set provides information on the fate of passengers on the fatal maiden voyage of the ocean liner Titanic, summarized according to economic status (class), sex, age and survival.

The dependent variable in this dataset is: survived(Categorical)

The independent Continous variables are: age and fare. <br>
The independent Categorical variables are: pclass, sex, sibsp, parch, embarked and alone.



```python
col_names = ['pclass', 'sex', 'age', 'sibsp', 'parch', 'fare',
             'embarked', 'alone', 'survived']

titanic_data = sns.load_dataset("titanic")

titanic_data = titanic_data[col_names]
titanic_data.head()

```

## Problem 1

Write a function for data preprocessing which will involve removing NaNs and converting categorical features into corresponding binarized features. Function should return a dataframe which should contain binarized features as well as numerical features along with the dependent variable.

<!-- grade_id: problem_1_answer, points: 0, grade: false, solution: true -->
```python
def preprocess_data(data):
    '''
    1. Preprocess data by removing all NaNs. 
    2. Create binarized features dataframe for each of the categorical features(can use get_dummies() function)
    3. Construct a dataframe for numerical feature
    4. Combine two DataFrames into a new DataFrame
    Parameters
    ----------
    data: dataframe containing titanic dataset
    
    Returns
    -------
    a dataframe 
    '''
    
    # YOUR CODE HERE

```

#### Sample output of the 1st 5 rows for the correct solution:

```python
data = preprocess_data(titanic_data)
data.head()

```

<!-- grade_id: problem_1_tests, points: 5, grade: true, solution: false -->
```python
assert_equal(len(data.index), 712)
assert_equal(len(data.columns), 12)

```

# Problem 2

We will now be using the above Titanic dataset to create a Decision tree model in order to predict whether a person survived(which is represented by binary 1) or died(which is represented by binary 0) based on the parameters provided in the data set.

In the code cell below do the following:

- Create a function for creating a Decision tree Classification model using sci-kit learn keeping random_state=40 and maximum depth as a parameter.<br>
- Fit your model on the training features and labels (which are stored in d_train and l_train).


```python
independent_vars = list(data)
independent_vars.remove('survived')
dependent_var = 'survived'

frac = 0.3
d_train, d_test, l_train, l_test = \
    train_test_split(data[independent_vars], data[dependent_var],
                    test_size=frac, random_state=40)

```

<!-- grade_id: problem_2_answer, points: 0, grade: false, solution: true -->
```python
def DecisionTree(data,tree_depth):
    '''
    
    Create a decision Tree Classification model using random_state=40 and max_depth as a parameters. 

    Parameters
    ----------
    data: dataframe containing titanic dataset
    tree_depth: max depth of the tree
    
    Returns
    -------
    A DecisionTree Classifier model 
    '''

    # YOUR CODE HERE

```

As the depth of the tree increases, the chance of tree overfitting on training dataset increases.

```python
titanic_model_3 = DecisionTree(data,tree_depth=3)
titanic_model_8 = DecisionTree(data,tree_depth=8)

predicted_3 = titanic_model_3.predict(d_test)
score_test_3 = 100.0 * metrics.accuracy_score(l_test, predicted_3)
print(f'Decision Tree Classification [Titanic Test Data(depth=3)] Score = {score_test_3:4.1f}%\n')

predicted_3 = titanic_model_3.predict(d_train)
score_train_3 = 100.0 * metrics.accuracy_score(l_train, predicted_3)
print(f'Decision Tree Classification [Titanic Train Data(depth=3)] Score = {score_train_3:4.1f}%\n')

predicted_8 = titanic_model_8.predict(d_test)
score_test_8 = 100.0 * metrics.accuracy_score(l_test, predicted_8)
print(f'Decision Tree Classification [Titanic Test Data(depth=8)] Score = {score_test_8:4.1f}%\n')

predicted_8 = titanic_model_8.predict(d_train)
score_train_8 = 100.0 * metrics.accuracy_score(l_train, predicted_8)
print(f'Decision Tree Classification [Titanic Train Data(depth=8)] Score = {score_train_8:4.1f}%\n')

```

<!-- grade_id: problem_2_tests, points: 5, grade: true, solution: false -->
```python
assert_equal(isinstance(titanic_model_3, DecisionTreeClassifier), True)
assert_equal(titanic_model_3.max_depth, 3)

```

## Problem 3

It is necessary to know what features gives us the most information or are the most important ones to create a model. Create a function to find the top 4 most important feature names along with the correspending value of importance.

<!-- grade_id: problem_3_answer, points: 0, grade: false, solution: true -->
```python
def importance(feature_names,model):
    '''
    Find top 4 most important feature names along with the correspending value of importance.  
    
    Parameters
    ----------
    feature_names : containing a list of independent variables(independent_vars)
    model : model name for the decision tree regressor
    
    Returns
    -------
    A sorted list of 4 elements in decending order containing a tuple of feature and importance as an element
    
    Example of what return output will look like
    --------------------------------------------
    [('sex_female', 0.43425212116045792),
     ('fare', 0.17077252568364368),
     ('pclass', 0.15921755986638889),
     ('age', 0.15394111874015437)]
    
    Hint
    ----
    Use zip(feature_names, model.feature_importances_) and use reversed and sorted functions to get the zip in sorted order. 
    Then convert this into a list and take the first 4 elements. 
    Or if you are not using reversed, then return the last 4 elements.
    '''
    
    # YOUR CODE HERE

```

<!-- grade_id: problem_3_tests, points: 10, grade: true, solution: false -->
```python
imp = importance(independent_vars,titanic_model_3)
a=[]
for i in range(len(imp)):
    a.append(np.round(imp[i][1],4))
assert_almost_equal(a,[0.6597, 0.2367, 0.1036, 0.0],places=4)
b=[]
for i in range(len(imp)):
    b.append(imp[i][0])

assert_equal(b,['sex_female', 'pclass', 'age', 'embarked_S'])

#Feature importance for titanic_model_3 will have the values : 
for i in range(len(imp)):
    print(f'{imp[i][0]} importance = {100.0*imp[i][1]:5.2f}%')


```

## Problem 4

We will be using Hitters dataset which contains player information including his performance and salary for Major League Baseball from the 1986 and 1987 seasons. This dataset was taken from the StatLib library which is maintained at Carnegie Mellon University. 

We will try to predict the Salary of players based on the factors available to us(information like Number of Hits, Number of Home Runs, etc.) in the dataset.



```python
wage = pd.read_csv('/home/data_scientist/data/misc/wages.csv')
wage['log_Salary'] = np.log(wage['Salary'])

```

```python
import patsy as pts 

y, x = pts.dmatrices('Salary + log_Salary ~ AtBat + Hits + HmRun + Runs + RBI + Walks + Years +' +
                     'CAtBat + CHits + CHmRun + CRuns + CRBI + CWalks + C(League) +' + 
                     'C(Division) + PutOuts + Assists + Assists + C(NewLeague)' ,
                     data=wage, return_type='dataframe')

```

In the code cell below do the following:

- Split data intro training:testing data set using random_state=23.
- Create 2 Decision tree models for Salary and log_Salary repectively using sci-kit learn keeping random_state=23, mae as error criterion and maximum features as 7.<br>
- Fit your model on the training features and labels (which are stored in ind_train and dep_train).
- Find the RMSE values from both the models for Salary.

<!-- grade_id: problem_4_answer, points: 0, grade: false, solution: true -->
```python
'''

- Split data intro training:testing data set(random_state=23)

- Create following 2 Decision Tree Models with:

1. a model named wage_model using Salary as dependent variable
2. a model named log_wage_model using log_Salary as dependent variable

- Using the above models, find :
1. the rmse value using wage_model: rmse_wage
2. the rmse value using log_wage_model: rmse_logwage

'''
frac = 0.3

# YOUR CODE HERE

```

<!-- grade_id: problem_4_tests, points: 10, grade: true, solution: false -->
```python
assert_true(type(wage_model), type(DecisionTreeRegressor))
assert_true(type(log_wage_model), type(DecisionTreeRegressor))
assert_almost_equal(423.44, round(rmse_logwage,2), places = 3)


```