# Week 4 Problem 3

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do not write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select *Kernel*, and restart the kernel and run all cells (*Restart & Run all*).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select *File* → *Save and CheckPoint*)

5. When you are ready to submit your assignment, go to *Dashboard* → *Assignments* and click the *Submit* button. Your work is not submitted until you click *Submit*.

6. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

7. **If your code does not pass the unit tests, it will not pass the autograder.**

# Due Date: 6 PM, February 12, 2018

```python
import pandas as pd
import numpy as np

import numpy.testing as npt
import matplotlib.pyplot as plt
from copy import deepcopy

from sklearn.linear_model import LinearRegression, Ridge, Lasso, ElasticNet
from sklearn.pipeline import make_pipeline as mp
from sklearn.preprocessing import PolynomialFeatures, Imputer
from sklearn.model_selection import train_test_split, GridSearchCV

from nose.tools import assert_equal, assert_almost_equal, assert_true
import warnings
warnings.filterwarnings("ignore")
```

The cell below reads in a simulated dataset where y (labels) is an unknown function of a,b, and c (your features).

```python
df = pd.read_csv('/home/data_scientist/data/misc/sim.reg.data')
df.head()
```

Output:
```
a         b           c         y
0  145.693523  0.226261    1.591351  2.512648
1   56.867985  0.002133    3.105082  1.582973
2  141.832455       NaN   58.017359  2.378522
3   19.585304  0.212010  172.831541  0.840794
4    0.131977  0.645755   88.145476  0.341427
```

# Problem 3.1

Complete the function *prob3_1* by performing a 70/30 split on the training and testing features and labels. Replace any of the missing values with the mean of the column. Return *(in this order)* the training features, the testing features, training labels, and testing labels.

<!-- grade_id: cell-26aa4513a0b76dc4, points: 0, grade: false, solution: true -->
```python
def prob3_1(data):
    '''
    Performs a splits on the features and labels. Replaces missing values with the mean of the column.
    
    Parameters
    ----------
    data: dataframe containing simulated dataset.
    
    Returns
    -------
    Training features as a multi dimensional numpy array (contains 70% of the features)
    Testing features as a multi dimensional numpy array (contains 30% of the features)
    Training labels as pandas series (contains 70% of the labels)
    Testing labels as pandas series (contains 30% of the labels)
    '''
    # YOUR CODE HERE
    x_train, x_test, y_train, y_test = \
        train_test_split(data[['a','b','c']], data['y'], test_size=0.3, random_state=0)
    
    imputer = Imputer(missing_values='NaN', strategy='mean',axis=0)
    x_train = imputer.fit_transform(x_train, y_train)
    x_test = imputer.fit_transform(x_test, y_test)


    
    return x_train, x_test, y_train, y_test

```

<!-- grade_id: cell-a324823d31c03f13, points: 10, grade: true, solution: false -->
```python
tX_train, tX_test, y_train, y_test = prob3_1(df.copy())

assert_equal(type(tX_train), type(np.ndarray(0)))
assert_equal(type(tX_test), type(np.ndarray(0)))
assert_equal(type(y_train), type(pd.core.series.Series()))
assert_equal(type(y_test), type(pd.core.series.Series()))


assert_almost_equal(tX_train[0][1], 0.4933, places=2)
assert_almost_equal(tX_train[20][1], 0.4933, places=2)


assert_almost_equal(tX_test[0][1], 0.5720, places=2)
assert_almost_equal(tX_test[20][1], 0.5362, places=2)


npt.assert_almost_equal(y_train.tolist()[:200], [2.212302679715205, 0.0812983126929061, 0.04727030279947797, 0.9829737188772332, 2.1686862194106054, 1.6123564685939362, 1.9485580592490084, 2.105117373729372, 2.1411761197498307, 1.6976345923763707, 1.5311313228785202, 2.162507027807703, 2.123054651382967, 1.2805432699567731, 2.4938123389877784, 0.6745831346046007, 3.5440260023585903, 0.04630999740476316, 1.7759938240532245, 2.0382880846699787, 1.1675949807638064, 1.6691116877873582, 0.039940353163266606, 2.286223325051137, 2.1192197218393023, 2.0055761748587693, -0.5477549988196424, 0.5947062609131056, 1.1426783548724595, 1.3298350523657616, 2.17072001172139, 2.183497206545365, 2.097181518029041, 0.1036257416416, 0.3677855252263231, 1.7541939527423105, 3.3045728395209086, 1.0063051087987138, 0.4896888131369209, 0.2696242613158241, 2.8112529448802728, 2.2358672220236238, -0.24768904743196296, 3.10353772893807, 1.9074634937713952, 1.3099938995993792, -0.4623774290130904, 1.7061960347983591, 0.4739988671363948, 0.4444170637955597, -0.05333685229962448, 1.148079640247882, 0.8357091605991601, 3.1947633490553526, 0.1150053544960732, 0.2980572927412253, 1.8954265397626864, 1.1470676973919958, 0.6746651155355947, 2.614187651741721, 1.2899019219138363, 1.9639662568666936, 1.7750985207858698, 1.522176106054589, 0.10601784522855184, 1.4517052794841887, 1.8799502416867848, 1.2310277021955829, 1.9976031947936743, 0.059376105076494164, 2.1741220862653217, -0.22677358181493024, 1.1280138671594038, 1.8981168446087149, 2.343093769786059, 0.9799402947266084, 0.2645921345400704, 1.0608214936885934, 2.163297891603712, 3.651033928034949, 2.3241005746167254, 1.5829727696817109, 0.4932984945537532, 1.964578553977544, 1.6082834944304083, 1.215035505817627, 2.4476853406319647, 1.4419602564907157, 1.708775723039362, 1.6549362700827812, 0.06232186355425218, 2.9478529430714118, 1.2563974955317891, 2.930222460025118, 1.254374571817709, 1.2581496312505405, 1.4204793222797405, 2.461025464522143, 0.011003146584178558, -0.0718395752936356, -0.4976949690756818, 1.2500424295004529, 0.9321311175148406, 2.388493234862328, 0.9719429650769642, -0.05113120701447815, 1.0945636277512731, 0.6219398712045754, 2.6200902148551433, 2.6440331491668787, 1.8612597493872516, 2.0553214921628564, 1.2818930640270616, 0.616271628636518, 0.6887409616849687, -0.4322759091098184, 0.09969344999675156, -0.4816557052902187, 1.1013979479897875, 0.8520645205491955, 1.8812568541916352, -0.4651486329050361, 0.7189695844272103, 0.6160963680593036, 0.10345306431747847, 0.7715187536325699, 2.324248732635162, 0.8417207317326361, 0.8052865605576498, 0.5199142785207596, 1.1267652564204758, 1.989251112852045, 1.2636650521657191, 2.358603329688961, -0.0912520203031717, 1.9610977209373708, 1.4195788382987204, 1.968138436983354, 1.7969886704825564, 0.2471217493931733, 1.9016364297537534, 1.3163547159479243, 2.0357934830650284, 1.9773481240881208, 1.2103374159629787, 2.052704013289523, 2.0096417342808013, 1.0267631818566565, 2.990665437527628, 1.9162142472806132, 1.2182818007525908, -0.4632280561505395, 1.835519248957408, 0.5189649705301519, 2.2167812208526327, -0.1787548860030752, 0.5773271104548295, 0.8941892611925739, 2.8627103530986955, 0.8372091270391041, 0.2708067619495871, 2.3238688487716668, 1.8570689910385072, 2.8568362217133685, 1.2056775443459695, 2.3578292500349423, 2.9516288724882167, 0.1980857881321452, 2.349465990021343, 1.0447899353656762, 0.5544109783162569, 1.3026846274618482, 1.9923427346579918, 0.4262273679099221, -0.0006704622116308068, 1.8110082033912975, 2.0755101732777548, 1.9142933285627504, 0.998083063956906, 2.903651075128166, 2.638588979682208, 1.7412272255980012, 0.1884054754534662, -0.14192356099644346, 1.3458933557009014, 0.21402032663751724, 2.16033852717896, 1.4459656638120304, 2.30940579729527, 1.2503252735053074, -0.09036718206799732, 1.4106602979777092, 0.37207339221924296, -0.2754578662142767, 1.2709498168780542, 1.7398413442981693, 0.4953424320880659, 3.0826778259078718, 1.0515807384625526, 1.7860358460899108], decimal=2)


```

# Problem 3.2

Complete the function *prob3_2* by fitting different order polynomial functions with [PolynomialFeatures](http://scikit-learn.org/stable/modules/generated/sklearn.preprocessing.PolynomialFeatures.html) on the training features and transform the training and testing features. *It's not necessary to search past the 15th order Polynomial function.* Using the transformed training features exhaustively search *(hint: use [GridSearchCV](http://scikit-learn.org/stable/modules/generated/sklearn.model_selection.GridSearchCV.html))* over the parameters for a [Ridge Regression](http://scikit-learn.org/stable/modules/generated/sklearn.linear_model.Ridge.html#sklearn.linear_model.Ridge) model passed in with *params* *( random_state of the Ridge Regression model to 0)*. Return the model that had the best CV score, along with transformed testing features by that polynominal function.

<!-- grade_id: cell-5f957f90aafce960, points: 0, grade: false, solution: true -->
```python
def prob3_2(tX_train, y_train, tX_test, params):
    '''
    Finds the best model.
    
    Parameters
    ----------
    tX_train - cleaned training features
    y_train - training labels
    tX_test - cleaned testing features
    params - parameter grid to be searched over
    
    Returns
    -------
    the best model found (this should be a GridSearchCV object)
    testing transformed feature matrix from the polynomial function that yielded the best CV score
    '''

    # YOUR CODE HERE
    best_cv = 0
    best_degree = 0
    for i in range(1,16):
        pp = PolynomialFeatures(i).fit(tX_train)
        tfx_train = pp.transform(tX_train)
        
        rr = Ridge(random_state=0)
        gsc = GridSearchCV(rr, param_grid=params)
        gsc.fit(tfx_train, y_train)
        if gsc.best_score_ > best_cv:
            best_cv = gsc.best_score_
            best_degree = i
            best_gsc = gsc
        
    best_pp = PolynomialFeatures(best_degree).fit(tX_train)
    Xte =  best_pp.transform(tX_test)
    
    return best_gsc, Xte
```

<!-- grade_id: cell-abb3d8eed2725005, points: 10, grade: true, solution: false -->
```python
search, Xte = prob3_2(tX_train, y_train, tX_test, params = {'alpha': [x/10 for x in range(0,20)]})
print("Best Model Parameter Found: ", search.best_params_)
print("Mean Cross Validation Score: ", search.best_score_)
print('-'*20)
assert_equal(type(search), type(GridSearchCV(None, {})), msg='Return a Grid Search')
assert_true(search.best_score_ >= .98)
```

Output:
```
Best Model Parameter Found:  {'alpha': 0.1}
Mean Cross Validation Score:  0.98882828499
--------------------
```

The cell below reads in a more complex simulated dataset, where y is some unknown function of all (or a subset of) the features a,b,c,d,e, and f.

```python
df2 = pd.read_csv('/home/data_scientist/data/misc/sim.reg.w4p3.data')

X = df2.drop('y', axis=1)
y = df2['y']

X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=.3, random_state=0)
df2.head()
```

# Problem 3.3

Complete the function *prob3_3* by creating different order polynomial order functions from the parameters *start* to *stop*. Fit the different [PolynomialFeature](http://scikit-learn.org/stable/modules/generated/sklearn.preprocessing.PolynomialFeatures.html)s on the training features and transform the training and testing features. Lastly perform Grid search on the parameters *params*  with an [ElasticNet model](http://scikit-learn.org/stable/modules/generated/sklearn.linear_model.ElasticNet.html) *(set the random_state argument to be 0)*. Return the best model, a list of the best cv_scores used with the ElasticNet model for each of the  polynomial combinations of the features created, and the transformed testing features of the model with the best CV score .

<!-- grade_id: cell-8173f3976a324cce, points: 0, grade: false, solution: true -->
```python
def prob3_3(X_train, y_train, X_test, start=1, stop=4, params = {'alpha': [x/10 for x in range(0,10)], 'l1_ratio': [x/10 for x in range(0,10)]}):
    
    '''
    Parameters
    ----------
    X_train - pandas dataframe containing training freatures
    y_train - pandas series containing training labels
    X_test - pandas dataframe containing testing features
    start - Polynomial degree to start at
    stop - Polynomial degree to stop at
    params - dictionary containing parameters to search over
    
    Returns
    -------
    Model with the best CV score (GridSearchCV object)
    CV Score of the best model (floating point number)
    Transformed Testing features (Multi dimensional numpy array)
    '''
    
    # YOUR CODE HERE
    best_cv = []
    best_degree = start
    
    for i in range(start, stop+1):
        pp = PolynomialFeatures(i).fit(X_train)
        tfx_train = pp.transform(X_train)
        tfx_test = pp.transform(X_test)
        
        enm = ElasticNet(random_state=0)
        gsc = GridSearchCV(enm, param_grid=params)
        gsc.fit(tfx_train, y_train)
        best_cv.append(gsc.best_score_)
        
        if gsc.best_score_ >= max(best_cv):
            best_degree = i
            best_gsc = gsc 
    
    best_pp = PolynomialFeatures(best_degree).fit(X_train)
    tfx_train = pp.transform(X_train)
    tfx_test = pp.transform(X_test)
    
    return gsc, best_cv, tfx_test
        
        
    
    
    
    
```

<!-- grade_id: cell-7b5e01bc90561d6c, points: 10, grade: true, solution: false -->
```python
model, cv_scores_3_4, Xte = prob3_3(X_train, y_train, X_test, start=3, stop=3)
print("Best CV Score from model:", model.best_score_)
npt.assert_almost_equal(cv_scores_3_4, [0.96360016442849883], decimal=2)

```

```python

```