# Week 3 Problem 3

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do not write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select *Kernel*, and restart the kernel and run all cells (*Restart & Run all*).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select *File* → *Save and CheckPoint*)

5. When you are ready to submit your assignment, go to *Dashboard* → *Assignments* and click the *Submit* button. Your work is not submitted until you click *Submit*.

6. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

7. **If your code does not pass the unit tests, it will not pass the autograder.**

# Due Date: 6 PM, February 05, 2018

```python
import pandas as pd
import numpy as np

from sklearn.model_selection import train_test_split
from sklearn.preprocessing import StandardScaler, Imputer, RobustScaler
from sklearn.pipeline import Pipeline

from sklearn.neighbors import KNeighborsClassifier
from sklearn.tree import DecisionTreeClassifier
from sklearn.svm import SVC
from sklearn.linear_model import LogisticRegression
from sklearn.ensemble import RandomForestClassifier, GradientBoostingClassifier, VotingClassifier

from nose.tools import assert_equal, assert_almost_equal, assert_true
```

In the cell below a simulated dataset is read in, where y (labels) is a unknown function of a, b, and c (your features).

```python
sim_data = pd.read_csv('/home/data_scientist/data/misc/sim.bNA.data')
sim_data.head()
```

Output:
```
a         b           c  y
0   51.367114  0.883546  376.457271  1
1  122.491822  0.309596   85.937127  2
2   25.628908  0.035058  296.998523  0
3   15.461822  0.797374  115.344490  1
4    0.661393  0.902865  131.885952  0
```

# Problem 3.1  

To complete problem 3.1 finish the function prob3_1 by performing a 70/30 split on the features and labels using [train_test_split](http://scikit-learn.org/stable/modules/generated/sklearn.model_selection.train_test_split.html) function *(Set the random_state argument to be 0)*. Then create a Pipeline that [Imputes](http://scikit-learn.org/stable/modules/generated/sklearn.preprocessing.Imputer.html) missing values and replaces them with the median of the column *(name this Imputer in your pipeline)*, and scales features that are robust to outliers *(name this Scaler in your Pipeline)*. Fit your pipeline on training features and labels, and transform the training and testing features. Finally return the pipeline, transformed training features, transformed testing features, training labels, and testing labels *(in this order)*.



<!-- grade_id: p3_1_answer, points: 0, grade: false, solution: true -->
```python
def prob3_1(data):
    '''
    Splits and cleans the simulated data passed in
    
    Parameters
    ----------
    data: pandas dataframe containing the simulated dataset
    
    Returns
    -------
    Pipeline containing a Imputer and Scaler
    Transformed training features as a multi dimensional numpy array
    Transformed testing features as a multi dimensional numpy array
    training labels as a pandas series
    testing labels as a pandas series
    '''
    # YOUR CODE HERE
    frac = 0.3
    x_train, x_test, y_train, y_test = \
        train_test_split(data[['a','b','c']], data['y'], test_size=frac, random_state=0)
    
    imp = Imputer(missing_values='NaN', strategy='median',axis=0,copy=False)
    rs = RobustScaler(copy=False)
    
    pipe = Pipeline([('Impute', imp), ('Scale', rs)])
    pipe = pipe.fit(x_train, y_train)
    new_x_train = pipe.transform(x_train)
    new_x_test = pipe.transform(x_test)
    
    return pipe, new_x_train, new_x_test, y_train, y_test

```

<!-- grade_id: p3_1_test, points: 10, grade: true, solution: false -->
```python
pipeline1, tX_train, tX_test, y_train, y_test = prob3_1(sim_data.copy())
assert_equal(type(pipeline1), type(Pipeline([('test', None)])), msg= 'you did not return a Pipeline')


assert_equal(y_train.tolist()[:100], [2, 2, 2, 0, 2, 0, 0, 1, 0, 1, 2, 2, 0, 0, 1, 1, 0, 0, 1, 2, 1, 0, 0, 1, 1, 1, 1, 3, 1, 1, 1, 1, 0, 0, 2, 0, 0, 2, 2, 0, 1, 0, 0, 2, 2, 1, 1, 0, 1, 2, 2, 2, 0, 1, 0, 2, 0, 1, 0, 2, 1, 1, 1, 1, 0, 1, 0, 0, 0, 1, 1, 0, 0, 1, 1, 1, 2, 1, 0, 2, 2, 2, 0, 0, 1, 2, 1, 0, 1, 0, 0, 1, 1, 2, 1, 2, 2, 0, 2, 3], msg="Check your pipeline and/or how you are splitting your data")
assert_equal(y_test.tolist()[:100], [0, 0, 1, 3, 0, 0, 0, 0, 1, 0, 2, 2, 0, 0, 2, 0, 1, 0, 0, 0, 1, 1, 0, 1, 2, 2, 0, 0, 1, 0, 2, 2, 1, 2, 3, 0, 0, 2, 1, 1, 0, 1, 0, 0, 1, 2, 2, 0, 1, 1, 1, 2, 1, 2, 0, 1, 1, 1, 2, 1, 0, 2, 0, 0, 1, 1, 1, 2, 1, 2, 1, 0, 2, 0, 0, 2, 1, 2, 1, 2, 0, 1, 0, 2, 0, 3, 0, 2, 1, 0, 3, 0, 0, 1, 3, 0, 2, 2, 0, 0], msg="Check your pipeline and/or how you are splitting your data")


assert_almost_equal(tX_train.tolist()[0][0], 0.28393067626221147, places=2)
assert_almost_equal(tX_train.tolist()[1][1], 0.5210341087236174, places=2)
assert_almost_equal(tX_train.tolist()[2][2], 1.2338664499983505, places=2)

assert_almost_equal(tX_test.tolist()[100][0], -0.18053703783351555, places=2)
assert_almost_equal(tX_test.tolist()[100][1], -0.6257219599565056, places=2)
assert_almost_equal(tX_test.tolist()[100][2], 0.27423619668613075, places=2)

```

# Problem 3.2
For Problem 3.2 compelete the function prob3_2 by creating a Voting Classifier that uses the following estimators: [GradientBoostingClassifier](http://scikit-learn.org/stable/modules/generated/sklearn.ensemble.GradientBoostingClassifier.html), [DecisionTreeClassifier](http://scikit-learn.org/stable/modules/generated/sklearn.tree.DecisionTreeClassifier.html), [LogisticRegression](http://scikit-learn.org/stable/modules/generated/sklearn.linear_model.LogisticRegression.html), [KNeighborsClassifier](http://scikit-learn.org/stable/modules/generated/sklearn.neighbors.KNeighborsClassifier.html)
, and a [RandomForestClassifier](http://scikit-learn.org/stable/modules/generated/sklearn.ensemble.RandomForestClassifier.html) (set the random_state argument to be 0 for these models). In your voting estimator your classifiers should respectively have the following names *gbc*, *dtc*, *lr*, *knn*, and *rfc*. Fit the estimators on the training data and return the voting classifier.

<!-- grade_id: p3_2_answer, points: 0, grade: false, solution: true -->
```python
def prob3_2(tX_train, y_train):
    '''
    Creates and fits a soft voting classifier using 5 unfitted classifiers
    
    Parameters
    ----------
    tX_train - transformed training features from Problem 3.1
    y_train - training labels from Problem 3.1
    
    Returns
    -------
    Fitted Voting Classifier
    
    '''
    # YOUR CODE HERE
    gbc = GradientBoostingClassifier(random_state=0)
    dtc = DecisionTreeClassifier(random_state=0)
    lr = LogisticRegression(random_state=0)
    knn = KNeighborsClassifier()
    rfc = RandomForestClassifier(random_state=0)
    
    ests = [('gbc', gbc), ('dtc', dtc), ('lr', lr), ('knn',knn), \
            ('rfc', rfc)]
    vclf = VotingClassifier(estimators=ests)
    vclf = vclf.fit(tX_train, y_train)
    
    return vclf

```

<!-- grade_id: cell-d6e71296bedc1831, points: 10, grade: true, solution: false -->
```python
vc = prob3_2(tX_train, y_train)
assert_equal(type(vc), type(VotingClassifier([('test', None)])))
score = vc.score(tX_test, y_test)
assert_true(score >= 0.93)


```

Reading in [Breast Cancer Dataset](http://archive.ics.uci.edu/ml/datasets/Breast+Cancer+Wisconsin+%28Original%29) from UCI's Machine Learning respository *([dataset description](http://archive.ics.uci.edu/ml/machine-learning-databases/breast-cancer-wisconsin/breast-cancer-wisconsin.names))*. The labels are contained the in *Class* column and the features are the other columns.

```python
names = ['id number', 'Clump Thickness', 'Uniformity Cell Size', 'Uniformity Cell Shape',
         'Marginal Adhesion', 'Single Epithelial Cell Size', 'Bare Nuclei', 'Bland Chromatin',
        'Normal Nucleoli', 'Mitoses', 'Class']
df = pd.read_csv('/home/data_scientist/data/misc/breast-cancer-wisconsin.data',
                 na_values='?', names=names, header=0)
df.head()
```

Output:
```
id number  Clump Thickness  Uniformity Cell Size  Uniformity Cell Shape  \
0    1000025                5                     1                      1   
1    1002945                5                     4                      4   
2    1015425                3                     1                      1   
3    1016277                6                     8                      8   
4    1017023                4                     1                      1   

   Marginal Adhesion  Single Epithelial Cell Size  Bare Nuclei  \
0                  1                            2            1   
1                  5                            7           10   
2                  1                            2            2   
3                  1                            3            4   
4                  3                            2            1   

   Bland Chromatin  Normal Nucleoli  Mitoses  Class  
0                3                1        1      2  
1                3           
```

# Problem 3.3  

For Problem 3.3 complete the function *prob3_3* by splitting the training and testing features/labels of the breast cancer data and creating 2 Pipelines that do the following:
- 1 - that replaces missings values with the most frequent value of the column (*name this Imputer*) and scales the data using a RobustScaler (*name this Scaler*)
- and 2- the other pipeline should be capabale of training and making predictions using a SVM (*name this svm*).  

Fit pipeline 1 on the training data and transform the training and testing features. Use the transformed training features and training labels to fit pipeline 2. Set the random_state argument to be 0 when appropriate. Return in the following order: *pipeline 1, pipeline 2, transformed testing features, and testing labels* 



<!-- grade_id: prob3_3_ans, points: 0, grade: false, solution: true -->
```python
def prob3_3(data):
    '''
    Creates 2 pipelines. The first Pipeline first removes missing values and scales the data.
    The second Pipeline trains a Support Vector Machine.
    
    Parameters
    ----------
    
    Returns
    -------
    Pipeline that removes missing values and scales the data
    Pipeline with a trained SVM
    Transformed testing features
    testing labels
    '''
    
    # YOUR CODE HERE
    frac = 0.7
    x_train, x_test, y_train, y_test = \
        train_test_split(data.iloc[:,:-1], data.iloc[:,-1], test_size=frac, random_state=0)
    
    imp = Imputer(missing_values='NaN', strategy='most_frequent', axis=0, copy=False)
    rs = RobustScaler(copy=False)
    
    pip1 = Pipeline([('Imputer', imp), ('Scaler', rs)])
    
    svm = SVC(random_state=0)
    pip2 = Pipeline([('svm', svm)])
    
    pip1.fit(x_train, y_train)
    new_x_train = pip1.transform(x_train)
    new_x_test = pip1.transform(x_test)
    
    pip2 = pip2.fit(new_x_train, y_train)
    
    
    return pip1, pip2, new_x_test, y_test
    

```

<!-- grade_id: prob3_3_test, points: 10, grade: true, solution: false -->
```python
pipeline2, pipeline3, tX_test, y_test = prob3_3(df.copy())
assert_equal(type(pipeline2), type(Pipeline([('test', None)])), msg='Make sure you return a Pipeline')

assert_equal(type(pipeline2.named_steps['Imputer']), type(Imputer()))

assert_equal(type(pipeline3), type(Pipeline([('test', None)])))
assert_equal(type(pipeline3.named_steps['svm']), type(SVC()))


score = pipeline3.score(tX_test, y_test)
assert_true(score >= 0.95)

```

```python

```